---
title       : Exploratorische Datenanalysen in R
description : Wie berechnet man exploratorische Faktorenanalysen in R? Wichtige Schritte sind dabei. 
--- type:NormalExercise lang:r xp:100 skills:1 key:8349a725f7
## <<<New Exercise>>>


*** =instructions
sad

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
