---
title       : Exploratorische und konfirmatorische Faktorenanalysen
description : Einige Grundlagen zum Thema
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf



--- type:NormalExercise lang:r xp:100 skills:1 key:7f0de9c3c4
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
